# Question 2

<!-- grade_id: 5c9323, points: 0, grade: false, solution: true -->
```python
def squares(n):  # modify function name and parameters
    """
    Compute the squares of the numbers from 1 to n.  [replace with function description]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: a702cf, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(squares(1), [1])
assert_equal(squares(2), [1, 4])

# Check that squares raises an error for invalid input:
from nose.tools import assert_raises
assert_raises(ValueError, squares, 0)
assert_raises(ValueError, squares, -1)

```